# Assignment 3
All questions are weighted the same in this assignment. This assignment requires more individual learning then the last one did - you are encouraged to check out the [pandas documentation](http://pandas.pydata.org/pandas-docs/stable/) to find functions or methods you might not have used yet, or ask questions on [Stack Overflow](http://stackoverflow.com/) and tag them as pandas and python related. All questions are worth the same number of points except question 1 which is worth 17% of the assignment grade.

**Note**: Questions 2-13 rely on your question 1 answer.

```python
import pandas as pd
import numpy as np

# Filter all warnings. If you would like to see the warnings, please comment the two lines below.
import warnings
warnings.filterwarnings('ignore')
```

### Question 1
Load the energy data from the file `assets/Energy Indicators.xls`, which is a list of indicators of [energy supply and renewable electricity production](assets/Energy%20Indicators.xls) from the [United Nations](http://unstats.un.org/unsd/environment/excel_file_tables/2013/Energy%20Indicators.xls) for the year 2013, and should be put into a DataFrame with the variable name of **Energy**.

Keep in mind that this is an Excel file, and not a comma separated values file. Also, make sure to exclude the footer and header information from the datafile. The first two columns are unneccessary, so you should get rid of them, and you should change the column labels so that the columns are:

`['Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable]`

Convert `Energy Supply` to gigajoules (**Note: there are 1,000,000 gigajoules in a petajoule**). For all countries which have missing data (e.g. data with "...") make sure this is reflected as `np.NaN` values.

Rename the following list of countries (for use in later questions):

```"Republic of Korea": "South Korea",
"United States of America": "United States",
"United Kingdom of Great Britain and Northern Ireland": "United Kingdom",
"China, Hong Kong Special Administrative Region": "Hong Kong"```

There are also several countries with numbers and/or parenthesis in their name. Be sure to remove these, e.g. `'Bolivia (Plurinational State of)'` should be `'Bolivia'`.  `'Switzerland17'` should be `'Switzerland'`.

Next, load the GDP data from the file `assets/world_bank.csv`, which is a csv containing countries' GDP from 1960 to 2015 from [World Bank](http://data.worldbank.org/indicator/NY.GDP.MKTP.CD). Call this DataFrame **GDP**. 

Make sure to skip the header, and rename the following list of countries:

```"Korea, Rep.": "South Korea", 
"Iran, Islamic Rep.": "Iran",
"Hong Kong SAR, China": "Hong Kong"```

Finally, load the [Sciamgo Journal and Country Rank data for Energy Engineering and Power Technology](http://www.scimagojr.com/countryrank.php?category=2102) from the file `assets/scimagojr-3.xlsx`, which ranks countries based on their journal contributions in the aforementioned area. Call this DataFrame **ScimEn**.

Join the three datasets: GDP, Energy, and ScimEn into a new dataset (using the intersection of country names). Use only the last 10 years (2006-2015) of GDP data and only the top 15 countries by Scimagojr 'Rank' (Rank 1 through 15). 

The index of this DataFrame should be the name of the country, and the columns should be ['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations',
       'Citations per document', 'H index', 'Energy Supply',
       'Energy Supply per Capita', '% Renewable', '2006', '2007', '2008',
       '2009', '2010', '2011', '2012', '2013', '2014', '2015'].

*This function should return a DataFrame with 20 columns and 15 entries, and the rows of the DataFrame should be sorted by "Rank".*

<!-- grade_id: cell-bce4d6f2ecdd1297, points: 0, grade: false, solution: true -->
```python
   

```

Output:
```
File "<ipython-input-20-6914a4df9c09>", line 31
    return merge_2.nsmallest(15,'Rank')
    ^
SyntaxError: 'return' outside function
```

```python
def answer_one():

    import pandas as pd
    import numpy as np

    energy = pd.read_excel('~/Documents/Course 1 - Notebook Resources/resources/week-3/datasets/Energy Indicators.xls',skiprows=9)
    energy=energy.drop(['Unnamed: 0','Unnamed: 1'],axis=1)
    energy=energy.drop([0],axis=0)
    energy=energy.rename(columns={'Country':'Country','Energy Supply':'Energy Supply','Energy Supply per capita':'Energy Supply per Capita','Renewable Electricity Production':'% Renewable'})
    energy['Energy Supply'] *=1000000
    energy=energy.replace('...',np.nan)

    energy['Country']=energy['Country'].str.replace(r"( \(.*\))","")
    energy['Country']=energy['Country'].str.replace(r"([0-9])","")

    energy=energy.replace('Republic of Korea','South Korea')                   #Method1
    energy.replace('United States of America','United States',inplace=True)    #Method2
    energy['Country'] = energy['Country'].replace({'China, Hong Kong Special Administrative Region':'Hong Kong','United Kingdom of Great Britain and Northern Ireland':'United Kingdom','Iran (Islamic Republic of)':'Iran'}) #Method3

    cmp=energy['Country']
    energy=energy[energy['Country']== cmp.values]
    energy.reset_index()

    ##############################################################################################
    GDP = pd.read_csv('~/Documents/Course 1 - Notebook Resources/resources/week-3/datasets/world_bank.csv', skiprows=4)
    GDP=GDP.replace('Korea, Rep.','South Korea')
    GDP=GDP.replace('Iran, Islamic Rep.','Iran')
    GDP=GDP.replace('Hong Kong SAR, China','Hong Kong')

    GDP [ GDP['Country Name']=='Hong Kong' ]
    GDP=GDP[['Country Name','2006','2007','2008','2009','2010','2011','2012','2013','2014','2015']]
    GDP=GDP.rename(columns={'Country Name':'Country'})
    ##############################################################################################

    ScimEn = pd.read_excel('~/Documents/Course 1 - Notebook Resources/resources/week-3/datasets/scimagojr-3.xlsx')
    ScimEn=ScimEn.nsmallest(15,'Rank')

    ##############################################################################################

    merge_1=pd.merge(energy,ScimEn,how='inner',left_on='Country',right_on='Country')
    final_merge=pd.merge(merge_1,GDP,on='Country')
    final_merge=final_merge[['Country','Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations', 'Citations per document', 'H index', 'Energy Supply', 'Energy Supply per Capita', '% Renewable', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015']]
    final_merge=final_merge.set_index('Country')


    return final_merge

answer_one()
```

Output:
```
Rank  Documents  Citable documents  Citations  \
Country                                                             
Australia             14       8831               8725      90765   
Brazil                15       8668               8596      60702   
Canada                 6      17899              17620     215003   
China                  1     127050             126767     597237   
France                 9      13153              12973     130632   
Germany                7      17027              16831     140566   
India                  8      15005              14841     128763   
Iran                  13       8896               8819      57470   
Italy                 11      10964              10794     111850   
Japan                  3      30504              30287     223024   
South Korea           10      11983              11923     114675   
Russian Federation     5      18534              18301      34266   
Spain                 12       9428               9330
```

<!-- grade_id: cell-780b5a4da845dbc3, points: 5, grade: true, solution: false -->
```python
assert type(answer_one()) == pd.DataFrame, "Q1: You should return a DataFrame!"

assert answer_one().shape == (15,20), "Q1: Your DataFrame should have 20 columns and 15 entries!"

```

<!-- grade_id: cell-74b5f0b971379f64, points: 10, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 2
The previous question joined three datasets then reduced this to just the top 15 entries. When you joined the datasets, but before you reduced this to the top 15 items, how many entries did you lose?

*This function should return a single number.*

```python
%%HTML
<svg width="800" height="300">
  <circle cx="150" cy="180" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="blue" />
  <circle cx="200" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="red" />
  <circle cx="100" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="green" />
  <line x1="150" y1="125" x2="300" y2="150" stroke="black" stroke-width="2" fill="black" stroke-dasharray="5,3"/>
  <text x="300" y="165" font-family="Verdana" font-size="35">Everything but this!</text>
</svg>
```

Output:
```
<IPython.core.display.HTML object>
```

<!-- grade_id: cell-b0c3202c000aada4, points: 0, grade: false, solution: true -->
```python
def answer_two():
        #Load xls file 
    xls_file = pd.ExcelFile('~/Documents/Course 1 - Notebook Resources/resources/week-3/datasets/Energy Indicators.xls')
    energy = xls_file.parse(skiprows=17,skip_footer=(38))
    energy = energy[['Unnamed: 1','Petajoules','Gigajoules','%']]
    # Rename the columns
    energy.columns = ['Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable']
    # Convert "..." to np.NaN
    energy[['Energy Supply', 'Energy Supply per Capita', '% Renewable']] =  energy[['Energy Supply', 'Energy Supply per Capita', '% Renewable']].replace('...',np.NaN).apply(pd.to_numeric)
    #Conevrt Energy Supply to gigajoules (there are 1,000,000 gigajoules in a petajoule)
    energy['Energy Supply'] = energy['Energy Supply']*1000000
    #Rename Countries
    energy['Country'] = energy['Country'].replace({"Republic of Korea": "South Korea", "United States of America": "United States",
                                                   "United Kingdom of Great Britain and Northern Ireland": "United Kingdom", 
                                                   "China, Hong Kong Special Administrative Region": "Hong Kong"})
    #Remove parantheses 
    energy['Country'] = energy['Country'].str.replace(r" \(.*\)","")
    
    #GDP read_csv & Skip header & only include 10 last years 
    GDP = pd.read_csv('~/Documents/Course 1 - Notebook Resources/resources/week-3/datasets/world_bank.csv', skiprows=4)
    #Rename Countries
    GDP = GDP[['Country Name','2006','2007','2008','2009','2010','2011','2012','2013','2014','2015']]
    GDP['Country Name'] = GDP['Country Name'].replace({"Korea, Rep.": "South Korea", "Iran, Islamic Rep.": "Iran", 
                                             "Hong Kong SAR, China": "Hong Kong"})
    #Only inlcude 10 last years of GDP 
    GDP.columns = ['Country','2006','2007','2008','2009','2010','2011','2012','2013','2014','2015']
    #Load scimagojr-3.xlsx
    ScimEn = pd.read_excel(io='~/Documents/Course 1 - Notebook Resources/resources/week-3/datasets/scimagojr-3.xlsx')
        
    #Join the three DataFrames on Country
    ScimEn_Energy = pd.merge(ScimEn, energy, how='inner', left_on='Country', right_on='Country')
    df_joint = pd.merge(ScimEn_Energy, GDP, how='inner', left_on='Country', right_on='Country')
    df_joint = df_joint.set_index('Country')
    df_joint_top = df_joint[:15]
    
    #The Actual answer is this, but (I guess) since the df_joint and df_joint_top were defined in the function answer_one() they 
    # were not defined here so I repeated all the code lines, But I think there should be something like 'global' to use the 
    # variables defined in other functions
    return len(df_joint.index) - len(df_joint_top.index)
answer_two()




```

Output:
```
147
```

<!-- grade_id: cell-be24cfcaa87ab071, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_two()) == int, "Q2: You should return an int number!"

```

### Question 3
What are the top 15 countries for average GDP over the last 10 years?

*This function should return a Series named `avgGDP` with 15 countries and their average GDP sorted in descending order.*

<!-- grade_id: cell-8c3d74335c0d489a, points: 0, grade: false, solution: true -->
```python
def answer_three():
    Top15 = answer_one()
    avgGDP = Top15[['2006','2007','2008','2009','2010','2011','2012','2013','2014','2015']].mean(axis=1).rename('avgGDP').sort_values(ascending=False)
    return avgGDP
answer_three()
```

Output:
```
Country
United States         1.536434e+13
China                 6.348609e+12
Japan                 5.542208e+12
Germany               3.493025e+12
France                2.681725e+12
United Kingdom        2.487907e+12
Brazil                2.189794e+12
Italy                 2.120175e+12
India                 1.769297e+12
Canada                1.660647e+12
Russian Federation    1.565459e+12
Spain                 1.418078e+12
Australia             1.164043e+12
South Korea           1.106715e+12
Iran                  4.441558e+11
Name: avgGDP, dtype: float64
```

<!-- grade_id: cell-aaaa11ef7d26f4cf, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_three()) == pd.Series, "Q3: You should return a Series!"

```

### Question 4
By how much had the GDP changed over the 10 year span for the country with the 6th largest average GDP?

*This function should return a single number.*

<!-- grade_id: cell-7f77d099e3e0bbee, points: 0, grade: false, solution: true -->
```python
def answer_four():
    Top15 = answer_one()
    avgGDP = answer_three()
    sixth_avg_GDP = avgGDP[avgGDP == avgGDP[5]]
    sixth_avg_GDP = sixth_avg_GDP.reset_index()
    Sixth_avgGDP_name = sixth_avg_GDP.loc[0][0]
    Top15.reset_index(inplace=True)
    diff = (Top15[Top15['Country'] == Sixth_avgGDP_name]['2015']-Top15[Top15['Country'] == Sixth_avgGDP_name]['2006']).astype(float)
    return diff.loc[3]
answer_four()
```

Output:
```
~/opt/anaconda3/lib/python3.8/site-packages/pandas/core/generic.py in xs(self, key, axis, level, drop_level)
   3489             loc, new_index = self.index.get_loc_level(key, drop_level=drop_level)
   3490         else:
-> 3491             loc = self.index.get_loc(key)
   3492 
   3493             if isinstance(loc, np.ndarray):

~/opt/anaconda3/lib/python3.8/site-packages/pandas/core/indexes/base.py in get_loc(self, key, method, tolerance)
   2895                 return self._engine.get_loc(casted_key)
   2896             except KeyError as err:
-> 2897                 raise KeyError(key) from err
   2898 
   2899         if tolerance is not None:

KeyError: 3
```

<!-- grade_id: cell-564dd8e5e24b0f83, points: 6.66, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 5
What is the mean energy supply per capita?

*This function should return a single number.*

<!-- grade_id: cell-58e79d558e982eef, points: 0, grade: false, solution: true -->
```python
def answer_five():
    Top15 = answer_one()
    mean = Top15["Energy Supply per Capita"].mean()
    return mean
answer_five()
```

Output:
```
157.6
```

<!-- grade_id: cell-30cc66180851638c, points: 6.66, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 6
What country has the maximum % Renewable and what is the percentage?

*This function should return a tuple with the name of the country and the percentage.*

<!-- grade_id: cell-b6824b78e74619f9, points: 0, grade: false, solution: true -->
```python
def answer_six():
    Top15 = answer_one()
    Max_Renewable = Top15[Top15['% Renewable'] == Top15['% Renewable'].max()] 
    return (Max_Renewable.index.tolist()[0],Max_Renewable['% Renewable'].tolist()[0])
answer_six()
```

Output:
```
('Brazil', 69.64803)
```

<!-- grade_id: cell-2bd201c5c7bdd80f, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_six()) == tuple, "Q6: You should return a tuple!"

assert type(answer_six()[0]) == str, "Q6: The first element in your result should be the name of the country!"

```

### Question 7
Create a new column that is the ratio of Self-Citations to Total Citations. 
What is the maximum value for this new column, and what country has the highest ratio?

*This function should return a tuple with the name of the country and the ratio.*

<!-- grade_id: cell-a4f39737f38aa53c, points: 0, grade: false, solution: true -->
```python
def answer_seven():
    Top15 = answer_one()
    Top15['Ratio_Self_Citation'] = Top15['Self-citations']/Top15['Citations']
    Max_Ratio_Self_Citation = Top15[Top15['Ratio_Self_Citation'] == Top15['Ratio_Self_Citation'].max()]
    return (Max_Ratio_Self_Citation.index.tolist()[0],Max_Ratio_Self_Citation['Ratio_Self_Citation'].tolist()[0])
answer_seven()
```

Output:
```
('China', 0.6893126179389422)
```

<!-- grade_id: cell-b7a163e9231b88c9, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_seven()) == tuple, "Q7: You should return a tuple!"

assert type(answer_seven()[0]) == str, "Q7: The first element in your result should be the name of the country!"

```

### Question 8

Create a column that estimates the population using Energy Supply and Energy Supply per capita. 
What is the third most populous country according to this estimate?

*This function should return the name of the country*

<!-- grade_id: cell-9ca58137846b84d6, points: 0, grade: false, solution: true -->
```python
def answer_eight():
    Top15 = answer_one()
    Top15['Population_Estimate'] = Top15['Energy Supply']/Top15['Energy Supply per Capita']
    df_Population = Top15['Population_Estimate'].sort_values(ascending=False).reset_index()
    Third_Most_Populous = df_Population['Country'][2]
    return Third_Most_Populous

answer_eight()
```

Output:
```
'United States'
```

<!-- grade_id: cell-3f3620c88df08b20, points: 0, grade: true, solution: false -->
```python
assert type(answer_eight()) == str, "Q8: You should return the name of the country!"

```

### Question 9
Create a column that estimates the number of citable documents per person. 
What is the correlation between the number of citable documents per capita and the energy supply per capita? Use the `.corr()` method, (Pearson's correlation).

*This function should return a single number.*

*(Optional: Use the built-in function `plot9()` to visualize the relationship between Energy Supply per Capita vs. Citable docs per Capita)*

<!-- grade_id: cell-033679ea456bfb9d, points: 0, grade: false, solution: true -->
```python
def answer_nine():
    Top15 = answer_one()
    Top15['Population_Estimate'] = Top15['Energy Supply']/Top15['Energy Supply per Capita']
    Top15['Citable_Docs'] = Top15['Citable documents']/Top15['Population_Estimate']
    #Top15[['Country','Population_Estimate', 'Citable_Docs']]
    correlation = Top15['Citable_Docs'].corr(Top15['Energy Supply per Capita'])
    return correlation
answer_nine()
```

Output:
```
<__array_function__ internals> in average(*args, **kwargs)

~/opt/anaconda3/lib/python3.8/site-packages/numpy/lib/function_base.py in average(a, axis, weights, returned)
    413 
    414     if returned:
--> 415         if scl.shape != avg.shape:
    416             scl = np.broadcast_to(scl, avg.shape).copy()
    417         return avg, scl

AttributeError: 'float' object has no attribute 'shape'
```

```python
def plot9():
    import matplotlib as plt
    %matplotlib inline
    
    Top15 = answer_one()
    Top15['PopEst'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
    Top15['Citable docs per Capita'] = Top15['Citable documents'] / Top15['PopEst']
    Top15.plot(x='Citable docs per Capita', y='Energy Supply per Capita', kind='scatter', xlim=[0, 0.0006])
```

<!-- grade_id: cell-3cb5c699065a4a20, points: 6.66, grade: true, solution: false -->
```python
assert answer_nine() >= -1. and answer_nine() <= 1., "Q9: A valid correlation should between -1 to 1!"

```

### Question 10
Create a new column with a 1 if the country's % Renewable value is at or above the median for all countries in the top 15, and a 0 if the country's % Renewable value is below the median.

*This function should return a series named `HighRenew` whose index is the country name sorted in ascending order of rank.*

<!-- grade_id: cell-0fdf60e64bf1a4f9, points: 0, grade: false, solution: true -->
```python
def answer_ten():
    Top15 = answer_one()
    median = Top15['% Renewable'].median()
    Top15['HighRenew'] = (Top15['% Renewable'] >= median).astype(int)
    Sorted = Top15.sort('Rank')
    return Sorted.HighRenew
answer_ten()
```

Output:
```
<ipython-input-46-60fbc884402d> in answer_ten()
      3     median = Top15['% Renewable'].median()
      4     Top15['HighRenew'] = (Top15['% Renewable'] >= median).astype(int)
----> 5     Sorted = Top15.sort('Rank')
      6     return Sorted.HighRenew
      7 answer_ten()

~/opt/anaconda3/lib/python3.8/site-packages/pandas/core/generic.py in __getattr__(self, name)
   5137             if self._info_axis._can_hold_identifiers_and_holds_name(name):
   5138                 return self[name]
-> 5139             return object.__getattribute__(self, name)
   5140 
   5141     def __setattr__(self, name: str, value) -> None:

AttributeError: 'DataFrame' object has no attribute 'sort'
```

<!-- grade_id: cell-b29a631fd9a7730f, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_ten()) == pd.Series, "Q10: You should return a Series!"

```

### Question 11
Use the following dictionary to group the Countries by Continent, then create a DataFrame that displays the sample size (the number of countries in each continent bin), and the sum, mean, and std deviation for the estimated population of each country.

```python
ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}
```

*This function should return a DataFrame with index named Continent `['Asia', 'Australia', 'Europe', 'North America', 'South America']` and columns `['size', 'sum', 'mean', 'std']`*

<!-- grade_id: cell-a5e0c0df27304f98, points: 0, grade: false, solution: true -->
```python

def answer_eleven():
    Top15 = answer_one()
    ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}
    Top15['Population_Estimate'] = (Top15['Energy Supply']/Top15['Energy Supply per Capita']).astype(float)
    Top15.reset_index(inplace=True)
    Top15['Continent'] = [ContinentDict[country] for country in Top15['Country']]
    df = Top15.set_index('Continent').groupby(level=0)['Population_Estimate'].agg({'size': np.size, 'sum': np.sum, 'mean': np.mean,'std': np.std})
    return df
    
answer_eleven()
```

Output:
```
~/opt/anaconda3/lib/python3.8/site-packages/pandas/core/groupby/generic.py in aggregate(self, func, engine, engine_kwargs, *args, **kwargs)
    244             # but not the class list / tuple itself.
    245             func = maybe_mangle_lambdas(func)
--> 246             ret = self._aggregate_multiple_funcs(func)
    247             if relabeling:
    248                 ret.columns = columns

~/opt/anaconda3/lib/python3.8/site-packages/pandas/core/groupby/generic.py in _aggregate_multiple_funcs(self, arg)
    290             # GH 15931
    291             if isinstance(self._selected_obj, Series):
--> 292                 raise SpecificationError("nested renamer is not supported")
    293 
    294             columns = list(arg.keys())

SpecificationError: nested renamer is not supported
```

<!-- grade_id: cell-18d1a07971b25743, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_eleven()) == pd.DataFrame, "Q11: You should return a DataFrame!"

assert answer_eleven().shape[0] == 5, "Q11: Wrong row numbers!"

assert answer_eleven().shape[1] == 4, "Q11: Wrong column numbers!"

```

Output:
```
~/opt/anaconda3/lib/python3.8/site-packages/pandas/core/groupby/generic.py in aggregate(self, func, engine, engine_kwargs, *args, **kwargs)
    244             # but not the class list / tuple itself.
    245             func = maybe_mangle_lambdas(func)
--> 246             ret = self._aggregate_multiple_funcs(func)
    247             if relabeling:
    248                 ret.columns = columns

~/opt/anaconda3/lib/python3.8/site-packages/pandas/core/groupby/generic.py in _aggregate_multiple_funcs(self, arg)
    290             # GH 15931
    291             if isinstance(self._selected_obj, Series):
--> 292                 raise SpecificationError("nested renamer is not supported")
    293 
    294             columns = list(arg.keys())

SpecificationError: nested renamer is not supported
```

### Question 12
Cut % Renewable into 5 bins. Group Top15 by the Continent, as well as these new % Renewable bins. How many countries are in each of these groups?

*This function should return a Series with a MultiIndex of `Continent`, then the bins for `% Renewable`. Do not include groups with no countries.*

<!-- grade_id: cell-2ecd9a4076abd8f0, points: 0, grade: false, solution: true -->
```python
def answer_twelve():
    Top15 = answer_one()
    ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}
    Top15.reset_index(inplace=True)
    Top15['Continent'] = [ContinentDict[country] for country in Top15['Country']]
    Top15['Renew_bins'] = pd.cut(Top15['% Renewable'],5)
    df = Top15.groupby(['Continent','Renew_bins']).size()
    return df
answer_twelve()

```

Output:
```
Continent      Renew_bins      
Asia           (2.212, 15.753]     4
               (15.753, 29.227]    1
               (29.227, 42.701]    0
               (42.701, 56.174]    0
               (56.174, 69.648]    0
Australia      (2.212, 15.753]     1
               (15.753, 29.227]    0
               (29.227, 42.701]    0
               (42.701, 56.174]    0
               (56.174, 69.648]    0
Europe         (2.212, 15.753]     1
               (15.753, 29.227]    3
               (29.227, 42.701]    2
               (42.701, 56.174]    0
               (56.174, 69.648]    0
North America  (2.212, 15.753]     1
               (15.753, 29.227]    0
               (29.227, 42.701]    0
               (42.701, 56.174]    0
               (56.174, 69.648]    1
South America  (2.212, 15.753]     0
               (15.753, 29.227]    0
               (29.227, 42.701]    0
               (42.701, 56.174]    0
               (56.174, 69.648]    1
dtype: int64
```

<!-- grade_id: cell-6c665602d6babab9, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_twelve()) == pd.Series, "Q12: You should return a Series!"

assert len(answer_twelve()) == 9, "Q12: Wrong result numbers!"

```

### Question 13
Convert the Population Estimate series to a string with thousands separator (using commas). Use all significant digits (do not round the results).

e.g. 12345678.90 -> 12,345,678.90

*This function should return a series `PopEst` whose index is the country name and whose values are the population estimate string*

<!-- grade_id: cell-58eb0ee0921d93fb, points: 0, grade: false, solution: true -->
```python
def answer_thirteen():
    Top15 = answer_one()
    Top15['Population_Estimate'] = (Top15['Energy Supply']/Top15['Energy Supply per Capita']).astype(float)
    ans = Top15['Population_Estimate'].apply(lambda x : '{:,}'.format(x))
    return ans
answer_thirteen()
```

Output:
```
Country
Australia              23,316,017.316017315
Brazil                 205,915,254.23728815
Canada                  35,239,864.86486486
China                 1,367,645,161.2903225
France                  63,837,349.39759036
Germany                 80,369,696.96969697
India                 1,276,730,769.2307692
Iran                    77,075,630.25210084
Italy                  59,908,256.880733944
Japan                  127,409,395.97315437
South Korea            49,805,429.864253394
Russian Federation            143,500,000.0
Spain                    46,443,396.2264151
United Kingdom         63,870,967.741935484
United States          317,615,384.61538464
Name: Population_Estimate, dtype: object
```

<!-- grade_id: cell-10fee7228cf973f6, points: 6.74, grade: true, solution: false -->
```python
assert type(answer_thirteen()) == pd.Series, "Q13: You should return a Series!"

assert len(answer_thirteen()) == 15, "Q13: Wrong result numbers!"

```

### Optional

Use the built in function `plot_optional()` to see an example visualization.

```python
def plot_optional():
    import matplotlib as plt
    %matplotlib inline
    Top15 = answer_one()
    ax = Top15.plot(x='Rank', y='% Renewable', kind='scatter', 
                    c=['#e41a1c','#377eb8','#e41a1c','#4daf4a','#4daf4a','#377eb8','#4daf4a','#e41a1c',
                       '#4daf4a','#e41a1c','#4daf4a','#4daf4a','#e41a1c','#dede00','#ff7f00'], 
                    xticks=range(1,16), s=6*Top15['2014']/10**10, alpha=.75, figsize=[16,6]);

    for i, txt in enumerate(Top15.index):
        ax.annotate(txt, [Top15['Rank'][i], Top15['% Renewable'][i]], ha='center')

    print("This is an example of a visualization that can be created to help understand the data. \
This is a bubble chart showing % Renewable vs. Rank. The size of the bubble corresponds to the countries' \
2014 GDP, and the color corresponds to the continent.")
```

```python
def plot_optional():
    import matplotlib as plt
    %matplotlib inline
    Top15 = answer_one()
    ax = Top15.plot(x='Rank', y='% Renewable', kind='scatter', 
                    c=['#e41a1c','#377eb8','#e41a1c','#4daf4a','#4daf4a','#377eb8','#4daf4a','#e41a1c',
                       '#4daf4a','#e41a1c','#4daf4a','#4daf4a','#e41a1c','#dede00','#ff7f00'], 
                    xticks=range(1,16), s=6*Top15['2014']/10**10, alpha=.75, figsize=[16,6]);

    for i, txt in enumerate(Top15.index):
        ax.annotate(txt, [Top15['Rank'][i], Top15['% Renewable'][i]], ha='center')

    print("This is an example of a visualization that can be created to help understand the data. \
This is a bubble chart showing % Renewable vs. Rank. The size of the bubble corresponds to the countries' \
2014 GDP, and the color corresponds to the continent.")

plot_optional()
```

Output:
```
File "<ipython-input-55-f1ebce269e62>", line 9
    
    ^
SyntaxError: invalid character in identifier
```

```python

```